Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Purpose of this assignment
This assignment is supposed to test you on the `print` fuction, `if` statements, `lists` and binary search.

# fizzbuzz
This question is asked so much in technical interviews and on intro CS problem sets that it has become a meme. It is important to know how to solve it so you can start making fun of it too!


Implement the function `fizzbuzz` that prints out the numbers 1 to 100 in the terminal, but if a number is a multiple of 3 you should print "fizz", if the number is a multiple of 5 you should print "buzz", and if the number is a multiple of both 3 and 5 you should print "fizzbuzz". It will look like this:
```
1
2
fizz
4
buzz
fizz
7
8
fizz
buzz
11
fizz
13
14
fizzbuzz
16
...
98
fizz
buzz
```

<!-- grade_id: cell-d98876d99dffa069, points: 0, grade: false, solution: true -->
```python
def fizzbuzz():
    """
    Prints numbers from 1 to 100 with the following rules:
    
    - For numbers that are multiples of 3, prints "fizz".
    - For numbers that are multiples of 5, prints "buzz".
    - For numbers that are multiples of both 3 and 5, prints "fizzbuzz".
    - For all other numbers, prints the number itself.
    
    The function outputs each result on a new line in the terminal.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9a26cdd615c504d2, points: 5, grade: true, solution: false -->
```python
# Test Cases

import io
import sys

# Helper function to capture print output
def capture_output(func, *args):
    captured_output = io.StringIO()
    sys.stdout = captured_output
    result = func(*args)
    sys.stdout = sys.__stdout__
    return result, captured_output.getvalue().strip()

# Verbose test case
result, output = capture_output(fizzbuzz)
printed_lines = output.split("\n")


assert len(printed_lines) == 100, "You did not print the correct number of times"

expected_lines = [(0, "1"), (1, "2"), (2, "fizz"), (3, "4"), (4, "buzz"), (5, "fizz"), (6, "7"), 
                  (7, "8"), (8, "fizz"), (9, "buzz"), (14, "fizzbuzz"), (29, "fizzbuzz"), (97, "98"), (98, "fizz"), (99, "buzz")]
error_found = False

for i, expected in expected_lines:
    try:
        actual = printed_lines[i]
        assert actual == expected, f"Line {i+1} is incorrect. Expected '{expected}', but got '{actual}'."
    except IndexError:
        print(f"Error: Line {i+1} is missing. Expected '{expected}'.")
        error_found = True
    except AssertionError as e:
        print(e)
        error_found = True

if not error_found:
    print("All test cases passed!")

```

# Polynomial Evaluation
To Start off, recall from math class that a polynomial is a function that adds together scaled powers of some variable, like x. For instance
$$
f(x) = x + 1
$$
$$
f(x) = x^2 - 10x + 5
$$
$$
f(x) = x^{100} - 400x^{50}
$$
are all polynomials.

We will represent the coefficients of a polynomial in a list with the highest degree coefficient at the last index of the list. For instance
```python
[1, 0, 2]
```
represents
$$
2 x^{2} + 0x + 1
$$
and 
```python
[1, 2, 3, 4, 5]
```
represents
$$
5 x^{4} + 4x^3 + 3x^2+2x+1
$$

Implement the `poly_evaluate` function which takes in a number and list and outputs the polynomial evaluated at that number. For instance
```python
>>> poly_evaluate(2, [1, 0, 2])
9
```
because
$$
2(2^2)+1=9
$$

<!-- grade_id: cell-ddd9c6676fd382e1, points: 0, grade: false, solution: true -->
```python
def poly_evaluate(x, coefficients):
    """
    Evaluate a polynomial at a given value of x.

    The polynomial is represented by a list of coefficients, where the last element 
    corresponds to the highest degree term. The function computes the value of the 
    polynomial by summing the scaled powers of x.

    Args:
        x (float): The value at which to evaluate the polynomial.
        coefficients (list[float]): A list of coefficients representing the polynomial, 
            with the highest degree coefficient at the last index.

    Returns:
        float: The result of evaluating the polynomial at the given value of x.

    Examples:
        >>> poly_evaluate(2, [1, 0, 2])
        9
        >>> poly_evaluate(3, [1, 2, 3, 4])
        142
        >>> poly_evaluate(0, [5, -3, 2])
        5
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-1edc12fb865aee19, points: 5, grade: true, solution: false -->
```python
# TEST CASES

def test_poly_evaluate():
    # Test Case 1: Simple polynomial [1, 0, 2] -> 2x^2 + 0x + 1
    assert poly_evaluate(2, [1, 0, 2]) == 9, "Test Case 1 Failed"

    # Test Case 2: Linear polynomial [1, 1] -> x + 1
    assert poly_evaluate(3, [1, 1]) == 4, "Test Case 2 Failed"

    # Test Case 3: Polynomial with higher degree [1, 2, 3, 4, 5] -> 5x^4 + 4x^3 + 3x^2 + 2x + 1
    assert poly_evaluate(1, [1, 2, 3, 4, 5]) == 15, "Test Case 3 Failed"
    assert poly_evaluate(0, [1, 2, 3, 4, 5]) == 1, "Test Case 3a Failed"

    # Test Case 4: Polynomial with all zero coefficients
    assert poly_evaluate(10, [0, 0, 0, 0]) == 0, "Test Case 4 Failed"

    # Test Case 5: Constant polynomial [5] -> 5
    assert poly_evaluate(100, [5]) == 5, "Test Case 5 Failed"

    # Test Case 6: Polynomial with negative coefficients [-1, 0, 2] -> 2x^2 + (-1)
    assert poly_evaluate(3, [-1, 0, 2]) == 17, "Test Case 6 Failed"

    # Test Case 7: Fractional coefficients [0.5, 0.25] -> 0.25x + 0.5
    assert poly_evaluate(4, [0.5, 0.25]) == 1.5, "Test Case 7 Failed"

    # Test Case 8: Polynomial evaluated at a negative x [-1, 0, 1] -> x^2 - 1
    assert poly_evaluate(-2, [-1, 0, 1]) == 3, "Test Case 8 Failed"

    print("All test cases passed!")

# Run the tests
test_poly_evaluate()

```

# Something went wrong
The _cube root_ of a number, _n_, is defined to be the real number, such that when the number is cubed, it equals _n_.
$$
(\sqrt[3]{n})^3 = n
$$
The function below is supposed to find a good approximation for the cube root of any given number, but something is going wrong. Can you find and fix the error(s) in the code?

```python
def cube_root(n):
    """
    Approximates the cube root of a given number using the bisection method.

    This function computes an approximation of the cube root of a number `n`
    with an error tolerance of 0.001. It handles both positive and negative
    input values, returning the cube root with the correct sign.

    Parameters:
        n (float): The number to approximate the cube root for.

    Returns:
        float: An approximation of the cube root of `n`.

    Example:
        >>> cube_root(8)
        2.0
        >>> cube_root(-27)
        -3.0
        >>> cube_root(0.5)
        0.7937
    """
    # Keep track of the sign of the input number
    if n != abs(n):
        sign = 1
    else:
        sign = -1

    #cube_root(-n) = -cube_root(n)
    n = abs(n)
    
    lower_bound = 0
    upper_bound = min(n, 1)

    # repeat while the lower_bound cubed is far away from n
    while abs(lower_bound ** 3 - n) < 0.001:
        middle = (lower_bound - upper_bound)/2
        if middle**3 - n < 0:
            upper_bound = middle
        else: 
            upper_bound = middle

    # lower_bound is now close to the actual cube root
    return sign*lower_bound
            
```

<!-- grade_id: cell-79f019e1c25169eb, points: 5, grade: true, solution: false -->
```python
# TEST CASES

def test_cube_root():
    # Test Case 1: Perfect cube (positive)
    result = cube_root(8)
    assert abs(result - 2.0) < 0.001, f"Failed: cube_root(8) = {result}, expected ~2.0"

    # Test Case 2: Perfect cube (negative)
    result = cube_root(-27)
    assert abs(result - (-3.0)) < 0.001, f"Failed: cube_root(-27) = {result}, expected ~-3.0"

    # Test Case 3: Small positive number
    result = cube_root(0.5)
    assert abs(result - 0.7937) < 0.001, f"Failed: cube_root(0.5) = {result}, expected ~0.7937"

    # Test Case 4: Small negative number
    result = cube_root(-0.125)
    assert abs(result - (-0.5)) < 0.001, f"Failed: cube_root(-0.125) = {result}, expected ~-0.5"

    # Test Case 5: Large positive number
    result = cube_root(1000)
    assert abs(result - 10.0) < 0.001, f"Failed: cube_root(1000) = {result}, expected ~10.0"

    # Test Case 6: Large negative number
    result = cube_root(-343000)
    assert abs(result - (-70.0)) < 0.001, f"Failed: cube_root(-343000) = {result}, expected ~-70.0"

    # Test Case 7: Zero
    result = cube_root(0)
    assert abs(result - 0.0) < 0.001, f"Failed: cube_root(0) = {result}, expected ~0.0"

    # Test Case 10: Non-perfect cube
    result = cube_root(15)
    expected = 15 ** (1/3)  # Use Python's built-in cube root approximation for comparison
    assert abs(result - expected) < 0.001, f"Failed: cube_root(15) = {result}, expected ~{expected}"

    print("All test cases passed!")

# Run the test cases
test_cube_root()

```

# Searching an ordered list
You are tasked with implementing an algorithm to determine if a given target element exists in a sorted list of integers. Implement the function `binary_search` that takes in the list `nums` and the integer `target` and returns `True` if `nums` is in `Target` and `False` otherwise.

Simply using `return target in nums` is not correct because this is not a binary search algorithm and runs in O(n) time instead of O(log(n)) time.

<!-- grade_id: cell-4049d323231415d3, points: 0, grade: false, solution: true -->
```python
def binary_search(nums, target):
    """
    Perform binary search to determine if a target exists in a sorted list.

    Args:
        nums (list[int]): A list of integers sorted in non-decreasing order.
        target (int): The integer value to search for in the list.

    Returns:
        bool: True if the target exists in the list, False otherwise.

    Example:
        >>> binary_search([1, 3, 5, 7, 9, 11], 5)
        True

        >>> binary_search([2, 4, 6, 8, 10], 3)
        False

        >>> binary_search([-10, -5, 0, 5, 10], -5)
        True

    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-3424cd9193d05c4b, points: 0, grade: true, solution: false -->
```python
# TEST CASES

assert binary_search([1, 3, 5, 7, 9, 11], 5) == True, "Got False; Expected True"
assert binary_search([1, 3, 5, 7, 9, 11], 11) == True, "Got False; Expected True"
assert binary_search([1, 3, 5, 7, 9, 11], 1) == True, "Got False; Expected True"
assert binary_search([1, 3, 5, 7, 9, 11], 60) == False, "Got True; Expected False"
assert binary_search(list(range(50, 200, 2)), 101) == False, "Got True; Expected False"
assert binary_search(list(range(200)), 200) == False, "Got True; Expected False"
assert binary_search(list(range(0, 100000, 3)), 999) == True, "Got False; Expected True"

print("All test cases passed!")
```

# More Practice

If you want to get some more practice with binary search and programatic problem solving skills, check out these leetcode challenges that you can solve in python. 

[Find peak element](https://leetcode.com/problems/find-peak-element/description/) 

[Find peak element II](https://leetcode.com/problems/find-a-peak-element-ii/description/) (binary search with 2D array)