Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Rule Generating
Recall from the reading that in an elementary cellular atomata we needed to make a rules list that represented the binary expansion of a number between 0 and 255. For example
```python
# Define the rule set as a list of binary values (e.g., Rule 220)
rule = [1, 1, 0, 1, 1, 1, 0, 0]
```
It is tedious to have to type this out for each rule that I want to test. Let's automate this process with a function that takes in the rule number and returns the rule list associated with it. For example
```python
>>> generate_rule_list(220)
[1, 1, 0, 1, 1, 1, 0, 0]
```
It may be helpful to use python's `bin` function that takes in an integer and spits out a string representing the number in binary.
```python
>>> bin(255)
'0b11111111'
>>> bin(0)
'0b0'
>>> bin(220)
'0b11011100'
```

<!-- grade_id: cell-ab327af703613ccf, points: 0, grade: false, solution: true -->
```python
def generate_rule_list(rule_number):
    """
    Generate a list of binary values representing the rule set for a given 
    elementary cellular automaton rule number.

    The rule number must be an integer between 0 and 255, inclusive. The output 
    is a list of 8 binary values corresponding to the binary representation of 
    the rule number, with the most significant bit (MSB) first.

    Args:
        rule_number (int): The rule number to convert, an integer between 0 and 255.

    Returns:
        list[int]: A list of 8 binary values (0s and 1s) representing the rule set.

    Examples:
        >>> generate_rule_list(220)
        [1, 1, 0, 1, 1, 1, 0, 0]
        >>> generate_rule_list(0)
        [0, 0, 0, 0, 0, 0, 0, 0]
        >>> generate_rule_list(255)
        [1, 1, 1, 1, 1, 1, 1, 1]
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-4cc6bd9a04bc080e, points: 5, grade: true, solution: false -->
```python
assert generate_rule_list(0) == [0, 0, 0, 0, 0, 0, 0, 0]
assert generate_rule_list(220) == [1, 1, 0, 1, 1, 1, 0, 0]
assert generate_rule_list(255) == [1, 1, 1, 1, 1, 1, 1, 1]
```

# Dead or Alive?
Now let's take a look at Conway's game of life that we discussed in lecture. Recall that this is a cellular automata played on a 2D grid of boxes. The state of a cell in the next generation is given by 4 rules dealing with the number of neighbors around a cell.


* Rule 1: Any alive cell surrounded by 1 or fewer alive cells dies by isolation
* Rule 2: Any alive cell surrounded by 2 or 3 alive cells survives to the next generation
* Rule 3: Any alive cell surrounded by 4 or more alive cells dies by overcrowding
* Rule 4: Any dead cell surrounded by exactly 3 alive cells becomes alive next generation through reproduction 


Lets implement a function called new_cell_state that takes in the current state of a cell, 1 for alive and 0 for dead, and the number of neighbors around the cell and returns the new state of the cell in the next generation.

<!-- grade_id: cell-7eeeaaa63c4899ed, points: 0, grade: false, solution: true -->
```python
def new_cell_state(current_state, neighbor_count):
    """
    Determine the next state of a cell in Conway's Game of Life based on its 
    current state and the number of alive neighbors.

    The function follows the rules of Conway's Game of Life:
    - Rule 1: Any alive cell (1) with 1 or fewer neighbors dies by isolation.
    - Rule 2: Any alive cell (1) with 2 or 3 neighbors survives.
    - Rule 3: Any alive cell (1) with 4 or more neighbors dies by overcrowding.
    - Rule 4: Any dead cell (0) with exactly 3 neighbors becomes alive.

    Args:
        current_state (int): The current state of the cell, 1 for alive and 0 for dead.
        neighbor_count (int): The number of alive neighbors around the cell.

    Returns:
        int: The new state of the cell in the next generation, 1 for alive and 0 for dead.

    Examples:
        >>> new_cell_state(1, 1)
        0
        >>> new_cell_state(1, 3)
        1
        >>> new_cell_state(0, 3)
        1
        >>> new_cell_state(0, 2)
        0
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-a4a4c666506c3ca0, points: 5, grade: true, solution: false -->
```python
assert new_cell_state(1, 1) == 0
assert new_cell_state(1, 3) == 1
assert new_cell_state(0, 3) == 1
assert new_cell_state(0, 2) == 0
```

# Hello, neighbor!
Now let's calculate the number of neighbors around a given cell in Conway's Game of Life. The game state will be represented as a 2D grid (a list of lists), where each cell is either 1 (alive) or 0 (dead). For example, valid game states might look like:

```python
[[0, 0, 1, 0],
 [1, 1, 0, 0],
 [0, 0, 0, 0],
 [0, 1, 1, 0]]
```
or

```python
[[0, 0, 1, 0, 1, 0],
 [1, 1, 0, 0, 0, 0],
 [0, 0, 0, 0, 1, 1],
 [0, 1, 1, 0, 1, 1]]
```

Write a function called get_neighbors that takes in two parameters:

1. game_state (list of lists of integers): The current state of the game grid.
2. coordinate (tuple of two integers): The (row, column) coordinate of a cell in the grid.

The function should return the number of alive neighbors (1s) around the given coordinate. Neighbors are the 8 cells directly adjacent (horizontally, vertically, and diagonally).

Special rules for edge cases:

If the coordinate is on the edge or corner of the grid, count only the neighbors that are within the bounds of the grid.
Treat all cells outside the grid as 0 (dead).
Example Usage:

```python
>>> game_state = [[0, 0, 1, 0],
                  [1, 1, 0, 0],
                  [0, 0, 0, 0],
                  [0, 1, 1, 0]]

>>> get_neighbors(game_state, (1, 1))  # Cell at row 1, column 1
3

>>> get_neighbors(game_state, (0, 2))  # Cell at row 0, column 2
1

>>> get_neighbors(game_state, (3, 3))  # Cell at row 3, column 3
2
Implement the get_neighbors function to calculate the correct number of neighbors for any valid coordinate.
```


<!-- grade_id: cell-bdf27036f8e0c011, points: 0, grade: false, solution: true -->
```python
def get_neighbors(game_state, coordinate):
    """
    Calculate the number of alive neighbors for a given cell in a game state grid.

    Args:
        game_state (list[list[int]]): The current game state represented as a 2D list of 0s (dead) and 1s (alive).
        coordinate (tuple[int, int]): A tuple representing the (row, column) of the cell.

    Returns:
        int: The number of alive neighbors (1s) around the given cell.

    Example:
        >>> game_state = [[0, 0, 1, 0],
                          [1, 1, 0, 0],
                          [0, 0, 0, 0],
                          [0, 1, 1, 0]]
        >>> get_neighbors(game_state, (1, 1))
        3
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-beeff244cbd25c5f, points: 5, grade: true, solution: false -->
```python
game_state = [
    [0, 0, 1, 0],
    [1, 1, 0, 0],
    [0, 0, 0, 0],
    [0, 1, 1, 0]
]
assert get_neighbors(game_state, (1, 1)) == 2

assert get_neighbors(game_state, (0, 2)) == 1

assert get_neighbors(game_state, (3, 3)) == 1

assert get_neighbors(game_state, (2, 2)) == 3

game_state = [
    [0, 0, 1, 0, 1, 0],
    [1, 1, 0, 0, 0, 0],
    [0, 0, 0, 0, 1, 1],
    [0, 1, 1, 0, 1, 1]
]
assert get_neighbors(game_state, (0, 4)) == 0

assert get_neighbors(game_state, (0, 0)) == 2

game_state = [
    [1, 1, 1],
    [1, 1, 1],
    [1, 1, 1]
]
assert get_neighbors(game_state, (1, 1)) == 8

game_state = [
    [0, 0, 0],
    [0, 0, 0],
    [0, 0, 0]
]
assert get_neighbors(game_state, (1, 1)) == 0

```

# Putting everything together
Now we are ready to tie everything together. Let's make a function called get_next_generation. This function will take in a game state and output the next generation of cells

Now we are ready to tie everything together! Let's create a function called get_next_generation.

This function should take in the current game state (a 2D list of integers representing the grid, where 1 means a cell is alive and 0 means it is dead) and output the next generation of cells based on the rules of Conway's Game of Life. You may find it helpful to use your previously implemented functions new_cell_state and get_neighbors.

<!-- grade_id: cell-ae4584eee9a63486, points: 0, grade: false, solution: true -->
```python
def get_next_generation(game_state):
    """
    Computes the next generation of a given game state based on the rules of Conway's Game of Life.

    Parameters:
        game_state (list of list of int): A 2D list representing the current game state, where
            each cell is either 1 (alive) or 0 (dead).

    Returns:
        list of list of int: A 2D list representing the next generation of the game state, where
            each cell is updated according to the rules of Conway's Game of Life.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fe5094a19bda7417, points: 10, grade: true, solution: false -->
```python
game_state = [[0, 0],
              [0, 0]]

new_game_state = get_next_generation(game_state)
assert new_game_state == [[0, 0],
                          [0, 0]]

game_state = [[1, 1],
              [1, 1]]

new_game_state = get_next_generation(game_state)
assert new_game_state == [[1, 1],
                          [1, 1]]

game_state = [[1, 0],
              [0, 0]]

new_game_state = get_next_generation(game_state)
assert new_game_state == [[0, 0],
                          [0, 0]]

game_state = [[0, 0, 0, 0, 0, 0], 
              [0, 0, 0, 0, 0, 0],
              [0, 0, 0, 0, 0, 0],
              [0, 0, 1, 1, 1, 0],
              [0, 0, 0, 0, 0, 0],
              [0, 0, 0, 0, 0, 0]]

new_game_state = get_next_generation(game_state)
assert new_game_state == [[0, 0, 0, 0, 0, 0], 
                          [0, 0, 0, 0, 0, 0],
                          [0, 0, 0, 1, 0, 0],
                          [0, 0, 0, 1, 0, 0],
                          [0, 0, 0, 1, 0, 0],
                          [0, 0, 0, 0, 0, 0]]

game_state = [[0, 0, 0, 0, 0, 1], 
              [0, 0, 1, 0, 1, 0],
              [0, 1, 0, 0, 0, 0],
              [0, 0, 1, 1, 1, 0],
              [0, 0, 1, 0, 0, 0],
              [0, 1, 0, 0, 1, 0]]

new_game_state = get_next_generation(game_state)
assert new_game_state == [[0, 0, 0, 0, 0, 0],
                          [0, 0, 0, 0, 0, 0],
                          [0, 1, 0, 0, 1, 0],
                          [0, 1, 1, 1, 0, 0],
                          [0, 1, 1, 0, 1, 0],
                          [0, 0, 0, 0, 0, 0]]
```

# You made it!
Congratulations on making it to the end! Feel free to play around with the code from here in your favorite IDE! See what happens if you change the rules. Can you come up with a game where the cells can be partially dead and partially alive (cell state is between 0 and 1)? What about making a probabilistic set of rules? Go nuts.