# All roads lead to...
We are given the (at least somewhat accurate) road_map of Italy as a dictionary. The keys are strings representing the names of Italian cities, and the values are sets of cities directly connected to each key city by road
```python
italy_road_map = {
    "Rome": {"Florence", "Naples", "Pescara", "Perugia"},
    "Florence": {"Rome", "Bologna", "Pisa"},
    "Naples": {"Rome", "Salerno", "Bari"},
    "Bologna": {"Florence", "Milan", "Venice"},
    "Milan": {"Bologna", "Turin", "Genoa", "Verona"},
    "Venice": {"Bologna", "Verona", "Trieste"},
    "Turin": {"Milan", "Genoa"},
    "Genoa": {"Milan", "Turin", "Pisa"},
    "Pisa": {"Florence", "Genoa"},
    "Verona": {"Milan", "Venice", "Trento"},
    "Trieste": {"Venice"},
    "Salerno": {"Naples", "Reggio Calabria"},
    "Bari": {"Naples", "Brindisi", "Taranto"},
    "Perugia": {"Rome", "Ancona"},
    "Ancona": {"Perugia", "Pescara"},
    "Pescara": {"Rome", "Ancona"},
    "Trento": {"Verona", "Bolzano"},
    "Bolzano": {"Trento"},
    "Reggio Calabria": {"Salerno"},
    "Brindisi": {"Bari"},
    "Taranto": {"Bari"}
}
```
Lets use our new DFS skills to make a simple navigation program. Given a starting city and ending city, I want to know the least number of cities I have to go through to get from the starting city to the ending city. If the path does not exist please return -1.


We will call this function `shortest_path`. For example
```
>>> shortest_path("Rome", "Florene", italy_road_map)
2
```
Because the path you take is `["Rome", "Florence"]`
```
>>> shortest_path(italy_road_map, "Rome", "Venice")
4
```
Because one path you can take is `['Rome', 'Florence', 'Bologna', 'Venice']`
```
>>> shortest_path(italy_road_map, "Rome", "Trieste")
5
```
Because one path you can take is `['Rome', 'Florence', 'Bologna', 'Venice', 'Trieste']`

<!-- grade_id: cell-5dad7f0cf741297d, points: 0, grade: false, solution: true -->
```python
def shortest_path(italy_road_map, start_city, end_city):
    """
    Find the shortest path between two cities in the Italy road map.

    Parameters:
        italy_road_map (dict): A dictionary representing road connections between cities.
        start_city (str): The city where the journey starts.
        end_city (str): The destination city.

    Returns:
        int: The least number of cities to go through to reach the destination,
             or -1 if no path exists.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6fff85eec718bd9c, points: 0, grade: true, solution: false -->
```python
# Test Cases

italy_road_map = {
    "Rome": {"Florence", "Naples", "Pescara", "Perugia"},
    "Florence": {"Rome", "Bologna", "Pisa"},
    "Naples": {"Rome", "Salerno", "Bari"},
    "Bologna": {"Florence", "Milan", "Venice"},
    "Milan": {"Bologna", "Turin", "Genoa", "Verona"},
    "Venice": {"Bologna", "Verona", "Trieste"},
    "Turin": {"Milan", "Genoa"},
    "Genoa": {"Milan", "Turin", "Pisa"},
    "Pisa": {"Florence", "Genoa"},
    "Verona": {"Milan", "Venice", "Trento"},
    "Trieste": {"Venice"},
    "Salerno": {"Naples", "Reggio Calabria"},
    "Bari": {"Naples", "Brindisi", "Taranto"},
    "Perugia": {"Rome", "Ancona"},
    "Ancona": {"Perugia", "Pescara"},
    "Pescara": {"Rome", "Ancona"},
    "Trento": {"Verona", "Bolzano"},
    "Bolzano": {"Trento"},
    "Reggio Calabria": {"Salerno"},
    "Brindisi": {"Bari"},
    "Taranto": {"Bari"}
}

assert shortest_path(italy_road_map, "Rome", "Venice") == 4

assert shortest_path(italy_road_map, "Rome", "Trieste") == 5

assert shortest_path(italy_road_map, "Rome", "Paris") == -1

assert shortest_path(italy_road_map, "Milan", "Trento") == 3

assert shortest_path(italy_road_map, "Bari", "Reggio Calabria") == 4
```

# A knightly puzzle
You have an n x m chess board with a single knight on it. Recall that in chess a knight moves in an L shape (2 squares in one direction and 1 square in the other direction). Implement the function `shortest_knight_moves` that takes in the dimensions of the board as a tuple (for example, the tuple `(6, 8)` corresponds to a board that is 6 squares wide by 8 squares tall), a tuple of the starting position of the knight (for example, `(1, 1)` corresponds to the knight starting in the upper left corner), and a tuple of the ending position of the knight and returns the length of the shortest sequence of valid moves to move the knight from the starting position to the ending position. If no such sequence of moves exists, return a path length of `-1`.

<!-- grade_id: cell-540f9fc2d774f9bf, points: 0, grade: false, solution: true -->
```python
def shortest_knight_moves(dimensions, start, end):
    """
    Find the shortest sequence of valid moves to move a knight from the start to the end position.
    
    Parameters:
        dimensions (tuple): A tuple (n, m) representing the width and height of the chessboard.
        start (tuple): A tuple (x, y) representing the starting position of the knight.
        end (tuple): A tuple (x, y) representing the ending position of the knight.
        
    Returns:
        int: The minimum number of moves needed to reach the end position from the start position,
             or -1 if no such sequence of moves exists.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-65b1e5e4d71b33d6, points: 0, grade: true, solution: false -->
```python
# Test Cases

assert shortest_knight_moves((10, 10), (1, 1), (1, 2)) == 3
# Because [(1, 1), (2, 3), (3, 1), (1, 2)]

assert shortest_knight_moves((10, 5), (1, 1), (2, 2)) == 4
# Because [(1, 1), (2, 3), (3, 5), (4, 3), (2, 2)]

assert shortest_knight_moves((8, 8), (8, 1), (8, 8)) == 5
# Because [(8, 1), (7, 3), (8, 5), (6, 4), (7, 6), (8, 8)]

assert shortest_knight_moves((8, 8), (1, 1), (1, 2)) == 3
# Because [(1, 1), (2, 3), (3, 1), (1, 2)]

assert shortest_knight_moves((1, 2), (1, 1), (1, 2)) == -1
# Because this is impossible

assert shortest_knight_moves((3, 2), (1, 1), (1, 2)) ==  -1
# Because this is impossible

assert shortest_knight_moves((100, 100), (50, 50), (100, 1)) == 33

assert shortest_knight_moves((100, 100), (50, 50), (51, 52)) == 1
# Because [(50, 50), (51, 52)]

assert shortest_knight_moves((8, 8), (4, 5), (4, 5)) == 0
```