Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Purpose of this assignment
This assignment is meant to help you get confortable with the Jupyter Notebook, practice the topics in lecture 0, and learn about if statements.

# Variable Assignment
Let's start off by recalling how to assign different types to variables. 


Set the variable `x` to the int 10, `y` to the float 0.5, `z` to the string "Italy", `my_list` to a list of `x`, `y`, and `z` in that order, and `my_tuple` to the tuple version of `my_list`.

<!-- grade_id: cell-a7ac898d18615c45, points: 0, grade: false, solution: true -->
```python
# Start with "x = " and proceed from there.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5ecc73f5751d0ae6, points: 5, grade: true, solution: false -->
```python
# Check if `x` is assigned correctly
assert 'x' in globals(), "Variable 'x' is not defined."
assert isinstance(x, int), f"Expected 'x' to be an int, but got {type(x).__name__}."
assert x == 10, f"Expected 'x' to be 10, but got {x}."

# Check if `y` is assigned correctly
assert 'y' in globals(), "Variable 'y' is not defined."
assert isinstance(y, float), f"Expected 'y' to be a float, but got {type(y).__name__}."
assert y == 0.5, f"Expected 'y' to be 0.5, but got {y}."

# Check if `z` is assigned correctly
assert 'z' in globals(), "Variable 'z' is not defined."
assert isinstance(z, str), f"Expected 'z' to be a string, but got {type(z).__name__}."
assert z == "Italy", f"Expected 'z' to be 'Italy', but got {z}."

# Check if `my_list` is assigned correctly
assert 'my_list' in globals(), "Variable 'my_list' is not defined."
assert isinstance(my_list, list), f"Expected 'my_list' to be a list, but got {type(my_list).__name__}."
assert my_list == [x, y, z], f"Expected 'my_list' to be [{x}, {y}, {z}], but got {my_list}."

# Check if `my_tuple` is assigned correctly
assert 'my_tuple' in globals(), "Variable 'my_list' is not defined."
assert isinstance(my_tuple, tuple), f"Expected 'my_list' to be a list, but got {type(my_tuple).__name__}."
assert my_tuple == (x, y, z), f"Expected 'my_list' to be [{x}, {y}, {z}], but got {my_tuple}."
```

# Copying and Aliasing

We talked about the dangers of list aliasing during the lecture. Make the variable `my_alised_list` that is aliased with `my_list` and the variable `my_copied_list` that is a copy of `my_list`.

<!-- grade_id: cell-9d40210a7542242c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3407c1307fcf5e4a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# AweSum function
Let's get some practice returning values from a function. Implement the function `add_numbers` that takes in two numbers and outputs their sum.

<!-- grade_id: cell-ec2e1b9c76bd9963, points: 0, grade: false, solution: true -->
```python
def add_numbers(num1, num2):
    """
    Adds two numbers and returns their sum.

    Parameters:
    num1 (int or float): The first number.
    num2 (int or float): The second number.

    Returns:
    int or float: The sum of num1 and num2.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-60b23e0e0c63d509, points: 5, grade: true, solution: false -->
```python
assert add_numbers(2, 2) == 4
assert add_numbers(2, -2) == 0
assert add_numbers(0, 0) == 0
assert add_numbers(7.2, -0.2) == 7
assert add_numbers(100, 100.34) == 200.34
```

# Verbose Calculator
Let's upgrade our `add_numbers`. Write a function `verbose_add_number`s that takes two numbers as arguments. The function should return their sum and also print the addition expression in the format:
`num1 + num2 = result`.

```python
>>> verbose_add_numbers(2, 2)
2 + 2 = 4 # This was printed
4 # This is the returned value
>>> verbose_add_numbers(30, 19)
30 + 19 = 49
49
>>> verbose_add_numbers(532, 982)
532 + 982 = 1514
1514
```
It may be helpful to know how to convert ints to strings:
```python
>>> my_int = 123
>>> type(my_int)
<class 'int'>
>>> my_str = str(my_int)
>>> type(my_str)
<class 'str'>
```
You could also take a look at the [python docs](https://docs.python.org/3/library/functions.html#print) for some cool features that the `print` function has.

<!-- grade_id: cell-e7e19ea37f013c03, points: 0, grade: false, solution: true -->
```python
def verbose_add_numbers(num1, num2):
    """
    Adds two numbers and prints the addition expression.

    Args:
        a (int or float): The first number.
        b (int or float): The second number.

    Returns:
        int or float: The sum of the two numbers.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-f189d395252824fe, points: 5, grade: true, solution: false -->
```python
import io
import sys

# Helper function to capture print output
def capture_output(func, *args):
    captured_output = io.StringIO()
    sys.stdout = captured_output
    result = func(*args)
    sys.stdout = sys.__stdout__
    return result, captured_output.getvalue().strip()

# Test 1
result, output = capture_output(verbose_add_numbers, 2, 2)
assert result == 4, f"Expected result to be 4, but got {result}."
assert output == "2 + 2 = 4", f"Expected output to be '2 + 2 = 4', but got '{output}'."

# Test 2
result, output = capture_output(verbose_add_numbers, 30, 19)
assert result == 49, f"Expected result to be 49, but got {result}."
assert output == "30 + 19 = 49", f"Expected output to be '30 + 19 = 49', but got '{output}'."

# Test 3
result, output = capture_output(verbose_add_numbers, 532, 982)
assert result == 1514, f"Expected result to be 1514, but got {result}."
assert output == "532 + 982 = 1514", f"Expected output to be '532 + 982 = 1514', but got '{output}'."

# Test 4: Negative numbers
result, output = capture_output(verbose_add_numbers, -5, 3)
assert result == -2, f"Expected result to be -2, but got {result}."
assert output == "-5 + 3 = -2", f"Expected output to be '-5 + 3 = -2', but got '{output}'."

# Test 5: Floats
result, output = capture_output(verbose_add_numbers, 3.5, 2.1)
assert result == 5.6, f"Expected result to be 5.6, but got {result}."
assert output == "3.5 + 2.1 = 5.6", f"Expected output to be '3.5 + 2.1 = 5.6', but got '{output}'."

print("All tests passed!")
```

# H4CK3R T3XT
Lets make a function that can up the coolness factor of any string by turning it into a h4ck3r string.

Replace all instances of 'e' or 'E' with '3'.
Replace all instances of 'a' or 'A' with '4'.
Replace all instances of 'i' or 'I' with '!'.
Replace all instances of 'o' or 'O' with '0'.
Replace all instances of 's' or 'S' with '$'.

For example:
```python
>>> to_h4ck3r_string("hacker")
h4ck3r
>>> to_h4ck3r_string("I love programming in Python!")
! l0v3 pr0gr4mm!ng !n Pyth0n!
```

You may want to check out some of Python's [builtin string methods](https://www.w3schools.com/python/python_ref_string.asp)

<!-- grade_id: cell-02787bc223f6004a, points: 0, grade: false, solution: true -->
```python
def to_h4ck3r_string(input_string):
    """
    Converts the input string into a 'h4ck3r' string by replacing specific characters.

    Parameters:
    input_string (str): The original string.

    Returns:
    str: The converted 'h4ck3r' string.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-99eacd74ff48f9da, points: 5, grade: true, solution: false -->
```python
# Example inputs and outputs
assert to_h4ck3r_string("hacker") == "h4ck3r"
assert to_h4ck3r_string("elite") == "3l!t3"
assert to_h4ck3r_string("I love programming in Python!") == "! l0v3 pr0gr4mm!ng !n Pyth0n!"
assert to_h4ck3r_string("PaSsWoRd") == "P4$$W0Rd"
```

# Conditional Execution


So far you may have thought that our code is a little bland. It executes in the same order every time. Let's change that with the `if` statement.


The `if` statement in python has a simple syntax structure.
```python
if CONDITION:
    CODE_TO_RUN_IF_CONDITION_IS_TRUE


CODE_AFTER_IF_STATEMENT
```
`CONDITION` is a boolean expression. A boolean expression is a piece of code that returns `True` or `False` when it is executed. Below are a couple of boolean expressions with numbers.


```python
>>> 3 == 3
True
>>> 3 == 4
False
>>> 3 < 4
True
>>> 3 > 4
False
```


For `a` and `b` numbers, Python gives us these Boolean operators to make Boolean expressions:
* Equals: `a == b` returns `True` if `a` and `b` have the same value
* Not Equals: `a != b`
* Less than: `a < b` returns `True` if a is strictly less than b
* Less than or equal to: `a <= b` returns `True` if `a` is less than `b` or `a` equals `b`
* Greater than: `a > b` returns `True` if a is strictly greater than `b`
* Greater than or equal to: `a >= b` returns `True` is `a` is greater than `b` or `a` equals `b`


If `CONDITION` returns `True` then `CODE_TO_RUN_IF_CONDITION_IS_TRUE` is executed. Othersie `CODE_TO_RUN` is skipped. Note that regaurdless of if `CONDITION` is `True` or `False`, the `CODE_AFTER_IF_STATEMENT` still runs. 
```python
x = 10


if x == 10: #check if x is equal to 10
    print("x has value 10")
    
print("DONE")
```
When you run this, both print statements indeed get executed. However, the program below only prints "DONE". 
```python
x = 11


if x == 10: #check if x is equal to 10
    print("x has value 10")
    
print("DONE")
```
This is because the boolean expression `x == 10` returned `False`.


If we want to execute another block of code in the case that our boolean expression returns `False`, we can use the `else` keyword.
```python
if CONDITION:
    CODE_TO_RUN_IF_CONDITION_IS_TRUE
else:
    CODE_TO_RUN_IF_CONDITION_IS_FALSE


CODE_AFTER_IF_STATEMENT
```
For example, 


```python
x = 10


if x == 10: #check if x is equal to 10
    print("x has value 10")
else:
    print("x does not have value 10")


print("DONE")
```


We can check for more conditions with the `elif` keyword (short for else-if).


```python
x = 10


if x > 10:
    print("x has value more than 10")
elif x == 10:
    print("x has value 10")
else:
    print("x has value less than 10")


print("DONE")
```
Note that once a boolean expression evaluates to `True` in a long list of `if` and `elif` statements, the rest of the `elif` statements are ignored. Take this code for example.
```python
x = 50


if x >= 100:
    print("x has value at least 100")
elif x >= 75:
    print("x has value at least 75")
elif x >= 50:
    print("x has value at least 50")
elif x >= 25:
    print("x has value at least 25")
else:
    print("x has value less than 25")


print("DONE")
```
When you run it only "x has value at least 50" is printed and "x has value at least 25" is not printed!


`if` statments will be very important for future assignments, so if you did not understand this short explanation I reccomend watching a few explaner vidoes on youTube.




# Even/Odd?
Let's put our new `if` skills to use. Write a function that takes in an integer and returns `"even"` if the integer is even and `"odd"` if the integer is odd.

<!-- grade_id: cell-d512500accec9025, points: 0, grade: false, solution: true -->
```python
def even_or_odd(n):
    """
    Determines if an integer is even or odd.

    Args:
        n (int): The integer to check.

    Returns:
        str: "even" if the integer is even, "odd" if the integer is odd.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-0f1aea875c8e1001, points: 5, grade: true, solution: false -->
```python
# Test 1: Even number
assert even_or_odd(2) == "even", "Test case failed: even_or_odd(2) should return 'even'."

# Test 2: Odd number
assert even_or_odd(3) == "odd", "Test case failed: even_or_odd(3) should return 'odd'."

# Test 3: Zero (even)
assert even_or_odd(0) == "even", "Test case failed: even_or_odd(0) should return 'even'."

# Test 4: Negative even number
assert even_or_odd(-8) == "even", "Test case failed: even_or_odd(-8) should return 'even'."

# Test 5: Negative odd number
assert even_or_odd(-7) == "odd", "Test case failed: even_or_odd(-7) should return 'odd'."

# Test 6: Large even number
assert even_or_odd(10**6) == "even", "Test case failed: even_or_odd(10**6) should return 'even'."

# Test 7: Large odd number
assert even_or_odd(10**6 + 1) == "odd", "Test case failed: even_or_odd(10**6 + 1) should return 'odd'."
```