Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# fizzbuzz
This question is asked so much in technical interviews and on intro CS problem sets that it has become a meme. However, it is still important to know how to solve it so you can start making fun of it too!

Implement the function `fizzbuzz` that prints out the numbers 1 to 100 in the terminal, but if a number is a multiple of 3 you should print "fizz", if the number is a multiple of 5 you should print "buzz", and if the number is a multiple of both 3 and 5 you should print "fizzbuzz". It will look like this: 
```
1
2
fizz
4
buzz
fizz
7
8
fizz
buzz
11
fizz
13
14
fizzbuzz
16
...
98
fizz
buzz
```

<!-- grade_id: cell-d98876d99dffa069, points: 0, grade: false, solution: true -->
```python
def fizzbuzz():
    """
    Prints the numbers from 1 to 100, but with the following exceptions:
    
    - For numbers that are multiples of 3, print "fizz" instead of the number.
    - For numbers that are multiples of 5, print "buzz" instead of the number.
    - For numbers that are multiples of both 3 and 5, print "fizzbuzz" instead of the number.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9a26cdd615c504d2, points: 5, grade: true, solution: false -->
```python
import io
import sys

# Helper function to capture print output
def capture_output(func, *args):
    captured_output = io.StringIO()
    sys.stdout = captured_output
    result = func(*args)
    sys.stdout = sys.__stdout__
    return result, captured_output.getvalue().strip()

result, output = capture_output(fizzbuzz)
printed_lines = output.split("\n")
assert printed_lines[0] == "1"
assert printed_lines[1] == "2"
assert printed_lines[2] == "fizz"
assert printed_lines[3] == "4"
assert printed_lines[4] == "buzz"
```

# Polynomial Evaluation (5 pts)
To Start off, recall from math class that a polynomial is a function that adds together scaled power of x. For instance
$$
f(x) = x + 1
$$
$$
f(x) = x^2 - 10x + 5
$$
$$
f(x) = x^{100} - 400x^{50}
$$
are all polynomials.

We will represent the coefficients of a polynomial in a list with the highest degree coefficient at the last index of the list. For instance
```python
[1, 0, 2]
```
represents
$$
2 x^{2} + 0x + 1
$$
and 
```python
[1, 2, 3, 4, 5]
```
represents
$$
5 x^{4} + 4x^3 + 3x^2+2x+1
$$

Implement the `poly_evaluate` function which takes in a number and list and outputs the polynomial evaluated at that number. For instance
```python
>>> poly_evaluate(2, [1, 0, 2])
9
```
because
$$
2(2^2)+1=9
$$

<!-- grade_id: cell-ddd9c6676fd382e1, points: 0, grade: false, solution: true -->
```python
def poly_evaluate(x, coefficients):
    """
    Evaluate a polynomial at a given value of x.

    The polynomial is represented by a list of coefficients, where the last element 
    corresponds to the highest degree term. The function computes the value of the 
    polynomial by summing the scaled powers of x.

    Args:
        x (float): The value at which to evaluate the polynomial.
        coefficients (list[float]): A list of coefficients representing the polynomial, 
            with the highest degree coefficient at the last index.

    Returns:
        float: The result of evaluating the polynomial at the given value of x.

    Examples:
        >>> poly_evaluate(2, [1, 0, 2])
        9
        >>> poly_evaluate(3, [1, 2, 3, 4])
        142
        >>> poly_evaluate(0, [5, -3, 2])
        5
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-1edc12fb865aee19, points: 5, grade: true, solution: false -->
```python
assert poly_evaluate(2, [1, 0, 2]) == 9
assert poly_evaluate(3, [1, 2, 3, 4]) == 142
assert poly_evaluate(0, [5, -3, 2]) == 5
assert poly_evaluate(1000, []) == 0
assert poly_evaluate(-10, [1, 0, -4.5, 1.2]) == -1649
```

# Something went wrong
The _cube root_ of a number, _n_, is defined to be the real number, such that when the number is cubed, it equals _n_.
$$
(\sqrt[3]{n})^3 = n
$$
The function below is supposed to find a good approximation for the cube root of any given number, but something is going wrong. Can you find the error(s) in the code?

<!-- grade_id: cell-135c917acfcba515, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-79f019e1c25169eb, points: 5, grade: true, solution: false -->
```python
from math import pow

assert abs(cube_root(0) - pow(0, 1/3)) < 0.001
assert abs(cube_root(1) - pow(1, 1/3)) < 0.001
assert abs(cube_root(-1) + pow(1, 1/3)) < 0.001
assert abs(cube_root(2) - pow(2, 1/3)) < 0.001
assert abs(cube_root(-2) + pow(2, 1/3)) < 0.001
assert abs(cube_root(100) - pow(100, 1/3)) < 0.001
```